# Pandas NBgrader exercises

In this notebook you will find the 6 deliverable exercises regarding pandas. This notebook is autograded so please dont modify any cell except to add your solution. Modifying this notebook may result in you not getting a grade.

```python
import pandas as pd
import numpy as np
```

```python
data = {'animal': ['cat', 'cat', 'snake', 'dog', 'dog', 'cat', 'snake', 'cat', 'dog', 'dog'],
        'age': [2.5, 3, 0.5, np.nan, 5, 2, 4.5, np.nan, 7, 3],
        'visits': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'priority': ['yes', 'yes', 'no', 'yes', 'no', 'no', 'no', 'yes', 'no', 'no']}

labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']


df = pd.DataFrame(data, index=labels)
```

**13.** Change the age in row 'f' to 1.5.

<!-- grade_id: cell-f8d4a64dd0fb6112, points: 0, grade: false, solution: true -->
```python
df.loc['f', 'age'] = 1.5
df
```

Output:
```
animal  age  visits priority
a    cat  2.5       1      yes
b    cat  3.0       3      yes
c  snake  0.5       2       no
d    dog  NaN       3      yes
e    dog  5.0       2       no
f    cat  1.5       3       no
g  snake  4.5       1       no
h    cat  NaN       1      yes
i    dog  7.0       2       no
j    dog  3.0       1       no
```

**18.** Sort `df` first by the values in the 'age' in *decending* order, then by the value in the 'visit' column in *ascending* order.

<!-- grade_id: cell-bb44e438722da485, points: 0, grade: false, solution: true -->
```python
df.sort_values(by=['age', 'visits'], ascending=[False, True])
```

Output:
```
animal  age  visits priority
i    dog  7.0       2       no
e    dog  5.0       2       no
g  snake  4.5       1       no
j    dog  3.0       1       no
b    cat  3.0       3      yes
a    cat  2.5       1      yes
f    cat  1.5       3       no
c  snake  0.5       2       no
h    cat  NaN       1      yes
d    dog  NaN       3      yes
```

**23.** Given a DataFrame of numeric values, say
```python
df = pd.DataFrame(np.random.random(size=(5, 3))) # a 5x3 frame of float values
```

how do you subtract the row mean from each element in the row?

<!-- grade_id: cell-5137285459670e80, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame(np.random.random(size=(5, 3)))
df.sub(df.mean(axis=1), axis=0)
```

Output:
```
0         1         2
0  0.011832 -0.413401  0.401569
1 -0.245279  0.073467  0.171812
2 -0.018378 -0.194724  0.213102
3  0.364202 -0.168064 -0.196138
4  0.017004 -0.264354  0.247350
```

**24.** Suppose you have DataFrame with 10 columns of real numbers, for example:

```python
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
```
Which column of numbers has the smallest sum? (Find that column's label.)

<!-- grade_id: cell-7830a1e1958310b9, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
print(df)
df.sum().idxmin()
```

Output:
```
a         b         c         d         e         f         g  \
0  0.431047  0.281623  0.156152  0.138227  0.706651  0.396365  0.324503   
1  0.907107  0.281674  0.123982  0.026719  0.995346  0.402933  0.743148   
2  0.606780  0.727255  0.684676  0.868137  0.555903  0.521356  0.267488   
3  0.984840  0.955091  0.920129  0.664171  0.013729  0.261575  0.701590   
4  0.673489  0.486187  0.457220  0.824876  0.013095  0.602819  0.962240   

          h         i         j  
0  0.680095  0.855406  0.719919  
1  0.541140  0.672882  0.249208  
2  0.442798  0.136413  0.548829  
3  0.591293  0.557722  0.667663  
4  0.814418  0.582987  0.172939  
'f'
```

The next three puzzles are slightly harder...

**26.** You have a DataFrame that consists of 10 columns of floating--point numbers. Suppose that exactly 5 entries in each row are NaN values. For each row of the DataFrame, find the *column* which contains the *third* NaN value.

(You should return a Series of column labels.)

<!-- grade_id: cell-e4f640a69590b406, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-182477a2beb50458, points: 3, grade: true, solution: false -->
```python
assert indexes == [3, 8, 5, 6, 8]
assert res.shape == (5,)
assert list(res.values) == [3, 8, 5, 6, 8]
```

**27.** A DataFrame has a column of groups 'grps' and and column of numbers 'vals'. For example: 

```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
```
For each *group*, find the sum of the three greatest values.

<!-- grade_id: cell-da30c3e27df05854, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})

df.groupby('grps')['vals'].nlargest(3).sum(level=0)
```

Output:
```
/var/folders/57/x7svmv1x253fpczm114w1chm0000gn/T/ipykernel_14173/1266516114.py:4: FutureWarning: Using the level keyword in DataFrame and Series aggregations is deprecated and will be removed in a future version. Use groupby instead. df.sum(level=1) should use df.groupby(level=1).sum().
  df.groupby('grps')['vals'].nlargest(3).sum(level=0)
grps
a    409
b    156
c    345
Name: vals, dtype: int64
```